---
title: 'Programming (in Python)'
description: 'What is programming, why would you (as a Physicist) want to learn it, and why are we teaching Python?'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Your very first Python program: learn to run code in this learning environment.

`@instructions`
Execute the code; you don't have to do anything else because we are just learning how to work with the Python window. Press the _Run Code_ button to see the output. Then press the _Submit Answer_ button.

`@hint`
Press the _Submit Answer_ button.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# Your very first Python program
print("Let's get started!")
```

`@solution`
```{python}
print("Let's get started!")
```

`@sct`
```{python}
Ex().has_printout(0)
```
